<!DOCTYPE html>
<html lang="en-US">
<!--********************************************-->
<!--*       Generated from PreTeXt source      *-->
<!--*                                          *-->
<!--*         https://pretextbook.org          *-->
<!--*                                          *-->
<!--********************************************-->
<head>
<meta http-equiv="Content-Type" content="text/html; charset=UTF-8">
<meta name="robots" content="noindex, nofollow">
</head>
<body class="ignore-math">
<article class="exercise exercise-like"><h3 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">5.1.10</span><span class="period">.</span><span class="space"> </span><span class="title">A Python program, including two others.</span>
</h3>
<div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="exercise-python-including-two">
<div class="ac_question col-md-12" id="exercise-python-including-two_question"><div class="para" id="p-810">Compute the total amount of money loaned and store it in the variable <code class="code-inline tex2jax_ignore">loan_total</code>.</div></div>
<textarea data-lang="python" data-timelimit="25000" data-audio="" data-coach="true" style="visibility: hidden;" id="exercise-python-including-two_editor" data-question_label="" data-codelens="true" data-include="hello-world statistics">loan_total = 0
for loan in loan_amount:
    loan_total += loan
print(loan_total)
====
from unittest.gui import TestCaseGui

class MyTests(TestCaseGui):

    def testOne(self):
        self.assertTrue('loan_total' in self.getEditorText(), "you need a loan_total variable")
        self.assertEqual(loan_total, sum(loan_amount), "Use the accumulator pattern to add up all the loans")
        self.assertFalse('sum(' in self.getEditorText(), "you may not use sum")

MyTests().main()
</textarea>
</div></div></div></article><span class="incontext"><a href="activecode.html#exercise-python-including-two" class="internal">in-context</a></span>
</body>
</html>
